# Practical Assignment — Centrality measures

```python
import networkx as nx
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd
```

In this assignment, we will use Moscow Metro underground stations taken from [Moscow Subway Optimal* Complete Route](https://medium.com/altsoph/moscow-subway-optimal-complete-route-ac98796dd652) research. Let us download and take a look at the data.

```python
station_url = 'https://raw.githubusercontent.com/vpozdnyakov/network_science_assignments/master/assignment_centrality_measures/stations_data.csv'
routes_url = 'https://raw.githubusercontent.com/vpozdnyakov/network_science_assignments/master/assignment_centrality_measures/routes_data.csv'
```

```python
stations = pd.read_csv(station_url)
routes = pd.read_csv(routes_url)
```

```python
stations.head()
```

```python
routes.head()
```

### Task 1. Metro graph (2 points)

Before we start analyzing centrality, let us build a graph of stations.

Write a function `metro_graph` that takes datasets `stations` and `routes` and returns a graph with all stations and routes of any type. Labels of nodes are in column `stations.id`. Nodes in `G.nodes` viewer should be in ascending order. All nodes should have attributes `station_name` and `line_name`.

<!-- grade_id: cell-96fa1082742a036a, points: 0, grade: false, solution: true -->
```python
def metro_graph(stations, routes):
    '''Build a graph with metro stations and routes'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c6f1237f3a4eaaa9, points: 0.7, grade: true, solution: false -->
```python
'''Check ordering of nodes, number of nodes, names of stations and lines'''
G = metro_graph(stations, routes)
assert (np.array(list(G.nodes)) == np.arange(194)).sum() == 194
assert G.number_of_edges() == 228
assert G.nodes[100]['station_name'] == 'Shabolovskaya'
assert G.nodes[100]['line_name'] == 'Kaluzhsko-Rizhskaya'
```

Let us draw the graph with default spring layout.

```python
plt.figure(figsize=(10, 8))
nx.draw(G, 
        width=0.5,
        node_size=60)
plt.show()
```

So far it does not look like a metro map. Let us specify a distance based layout using information from the dataset `stations`.

Write a function `metro_layout` that takes the dataset `stations` and returns a dictionary with nodes as keys and positions as values:
```
{node_id: [longitude, latitude]}
```

<!-- grade_id: cell-7ca2e105bb2493b1, points: 0, grade: false, solution: true -->
```python
def metro_layout(stations):
    '''Build a distance based layout of the metro graph'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-b5e28f603eaba3a6, points: 0.7, grade: true, solution: false -->
```python
'''Check a type and positions of some random nodes'''
pos = metro_layout(stations)
assert type(pos) == dict
assert len(pos) == 194
assert pos[100] == [37.6078018, 55.71882329]
assert pos[172] == [37.72840065, 55.69297728]
```

Let us draw the graph with the distance based layout.

```python
plt.figure(figsize=(10, 8))
nx.draw(G, 
        pos,
        width=0.5,
        node_size=60)
plt.show()
```

The final preprocessing step, let us colorize stations by their lines.

Write a function `station_color` that takes a graph `G` and returns a dictionary with nodes as keys and color as values. Colors are given in the `line_color` dictionary.

```python
line_color = {
    'Sokolnicheskaya': 'tab:red',
    'Zamoskvoretskaya': 'tab:green',
    'Arbatsko-Pokrovskaya': 'tab:blue',
    'Filevskaya': 'deepskyblue',
    'Koltsevaya': 'tab:brown',
    'Kaluzhsko-Rizhskaya': 'tab:orange',
    'Tagansko-Krasnopresnenskaya': 'tab:purple', 
    'Kalininskaya': 'gold',
    'Serpuhovsko-Timiryazevskaya': 'gray', 
    'Lyublinsko-Dmitrovskaya': 'yellowgreen',
    'Kahovskaya': 'paleturquoise', 
    'Butovskaya': 'lightblue', 
    'Monorels': 'royalblue'
}
```

<!-- grade_id: cell-2f1d4177661a9582, points: 0, grade: false, solution: true -->
```python
def station_color(G):
    '''Returns colors of stations'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-8777de6843545b61, points: 0.6, grade: true, solution: false -->
```python
'''Check a type and colors of some random nodes'''
nx.set_node_attributes(G, station_color(G), 'color')
assert type(station_color(G)) == dict
assert len(station_color(G)) == 194
assert station_color(G)[9] == 'tab:red'
assert station_color(G)[193] == 'royalblue'
assert station_color(G)[119] == 'tab:purple'
```

Let us draw the graph with the distance based layout and stations colors.

```python
plt.figure(figsize=(10, 8))
nx.draw(
    G, 
    pos,
    width=0.5,
    node_color=list(nx.get_node_attributes(G, 'color').values()),
    node_size=60,
)
plt.show()
```

### Task 2. Degree, closeness, betweenness centralities (2 points)

Let us look at centrality measures of the metro stations graph. First, use networkx to calculate degree, closeness, betweenness centralities.

Write a function `centralities` that takes a graph `G` and returns a dictionary with keys `degree`, `closeness`, `betweenness`. The values are np.arrays with values of each centrality.

<!-- grade_id: cell-8ef9147d12ad6ac4, points: 0, grade: false, solution: true -->
```python
def centralities(G):
    '''Calculate degree, closeness, betweenness centralities of the graph'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-dd093037fa7756f4, points: 1, grade: true, solution: false -->
```python
'''Check some random values of centralities'''
c = centralities(G)
assert c['degree'].shape[0] == 194
assert c['degree'][193] == 0.0051813471502590676
assert c['closeness'][94] == 0.12107904642409034
assert c['betweenness'][49] == 0.0991400402993669
```

Next, let us create a table with top 10 nodes for each centrality.

Write a function `top10_centralities` that takes a graph `G` and returns a dictionary with keys `degree`, `closeness`, `betweenness`. The values are pd.DataFrames with columns:
* `station` — station name
* `line` — line name
* `centrality` — value of a centrality
Each dataframe contains 10 rows with highest centrality values.

<!-- grade_id: cell-3dad4de0a9186415, points: 0, grade: false, solution: true -->
```python
def top10_centralities(G):
    '''Create table of top 10 nodes by degree, closeness, betweenness centralities'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0bfea33ff1fcf96c, points: 1, grade: true, solution: false -->
```python
'''Check that some random nodes are in the top 10 tables'''
top_dfs = top10_centralities(G)
assert 'Biblioteka imeni Lenina' in top_dfs['degree']['station'].to_list()
assert 'Serpuhovskaya' in top_dfs['betweenness']['station'].to_list()
assert 'Teatralnaya' in top_dfs['closeness']['station'].to_list()
assert 'Arbatsko-Pokrovskaya' in top_dfs['closeness']['line'].to_list()
```

Following code makes a visualization: for each centrality it prints a top 10 tables and graph, node size and color correspond to the centrality value.

```python
cases = [['degree', 5000, 0, 'Degree centrality'], 
         ['closeness', 700, 0, 'Closeness centrality'], 
         ['betweenness', 900, 20, 'Betweenness centrality']]
for c_key, scale, bias, title in cases:
    print(title)
    display(top_dfs[c_key])
    plt.figure(figsize=(10, 8))
    nx.draw(G, 
            pos,
            width=0.5,
            linewidths=0.5,
            edgecolors='black',
            cmap=plt.cm.hot,
            node_size=c[c_key]*scale + bias,
            node_color=c[c_key])
    plt.show()
```

### Task 3. Center of the graph (2 points)

In the next section we will slightly modify the initial graph — truncate all stations that located outside of Koltsevaya line.

Write a function `metro_graph_center` that takes a metro graph `G` and returns a subgraph with stations on Koltsevaya line and stations inside Koltsevaya line.

<!-- grade_id: cell-7e22abcc984c16d2, points: 0, grade: false, solution: true -->
```python
def metro_graph_center(G):
    '''Build a subgraph with stations inside Koltsevaya line'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-da2bb8cea6d1507c, points: 2, grade: true, solution: false -->
```python
'''Check the number of nodes and edges and some random stations'''
subG = metro_graph_center(G)
assert len(subG.nodes) == 57
assert len(subG.edges) == 85
assert 149 in list(subG.nodes)
assert 82 in list(subG.nodes)
assert 137 in list(subG.nodes)
```

Let us draw the center of the metro graph with station names and colors.

```python
plt.figure(figsize=(10, 8))
nx.draw(subG, 
        pos,
        width=0.5,
        node_color=list(nx.get_node_attributes(subG, 'color').values()),
        labels=nx.get_node_attributes(subG, 'station_name'),
        font_size=10)
plt.show()
```

### Task 4. Centralization (1 point)

The center of the metro graph looks a bit like a circular graph. We can see that a circular graph has the lowest centralization since

$$\text{Centralization} = \frac{\sum_i^N[C_x(p_*) - C_x(p_i)]}{\max\sum_i^N[C_x(p_*) - C_x(p_i)]}$$

where 
* $C_x$ — some centrality measure
* $p_i$ — $i$-th node
* $p_*$ — node with the largest centrality value
* $\max$ — maximum over all possible graphs with the same number of nodes (a star graph for degree, betweenness, closeness centralities).

Let us check how centralization changes during transformation our truncated graph to a pure circular graph.

Write a function `centralization_seq` that takes the truncated graph `G` and deletes stations inside Koltsevaya line one-by-one in ascending order. In each step the centralization is calculated. The function returns a dictionary with keys `degree`, `closeness`, `betweenness` and values — np.arrays with centralization values.

<!-- grade_id: cell-599d21555a10592f, points: 0, grade: false, solution: true -->
```python
def centralization_seq(G):
    '''Delete stations and calculate centralization for degree, closeness, betweenness'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-19b3b60a79fa0348, points: 1, grade: true, solution: false -->
```python
'''Check the output shape and some random centralization values'''
cent = centralization_seq(subG)
assert cent['degree'].shape == (45,)
assert cent['closeness'].shape == (45,)
assert cent['betweenness'].shape == (45,)
assert cent['degree'][-1] == cent['betweenness'][-1] == cent['closeness'][-1]
assert cent['degree'][41].round(4) == 0.0191
assert cent['degree'][2].round(4) == 0.0377
assert cent['closeness'][11].round(4) == 0.0604
assert cent['closeness'][9].round(4) == 0.0656
assert cent['betweenness'][14].round(4) == 0.0847
```

Let us draw how the centralization changes when the graph tends to a circular graph.

```python
plt.figure(figsize=(7, 5))
plt.plot(cent['degree'], label='Degree')
plt.plot(cent['closeness'], label='Closeness')
plt.plot(cent['betweenness'], label='Betweenness')
plt.ylabel('Centralization')
plt.xlabel('Number of deleted nodes')
plt.legend()
plt.show()
```

### Task 5. Katz centrality (2 points)

Katz centrality is based on the number of all paths (not only shortest) passing through a node. By the way, how to quickly calculate the total number of fixed length paths between two nodes?

Write a function `number_of_paths` that takes graph `G`, nodes `i`, `j` and the length of paths `n`. The function returns the number of paths.

<!-- grade_id: cell-99d5119d30c4e079, points: 0, grade: false, solution: true -->
```python
def number_of_paths(G, i, j, n):
    '''Calculate number of paths of length n'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0f40f6090975a9d5, points: 1, grade: true, solution: false -->
```python
'''Check some paths in the city center'''
# Paveletskaya (Koltsevaya) - Ohotny ryad. Length 4
# Single path: Paveletskaya (Koltsevaya) -> Paveletskaya (Zamoskvoretskaya) 
# -> Novokuznetskaya -> Teatralnaya -> Ohotny ryad
assert number_of_paths(subG, 78, 9, 4) == 1
# Novoslobodskaya - Pushkinskaya. Length 4
# 1st path: Novoslobodskaya -> Belorusskaya -> Krasnopresnenskaya
# -> Barrikadnaya -> Pushkinskaya
# 2nd path: Novoslobodskaya -> Mendeleevskaya -> Tsvetnoy bulvar
# -> Chehovskaya -> Pushkinskaya
assert number_of_paths(subG, 83, 120, 4) == 2
# Novoslobodskaya - Pushkinskaya. Length 7. Many paths.
assert number_of_paths(subG, 83, 120, 7) == 31
```

Now we can apply similar approach in calculation Katz centrality. The formula is

$$k_i = \beta\sum_j A_{ij} + \beta^2\sum_j A_{ij}^2 + \beta^3\sum_j A_{ij}^3 + \dots$$

where $\beta$ is an attenuation factor with restriction 

$$\beta < \frac{1}{\lambda}$$

where $\lambda$ is a largest eigenvalue of a matrix $A$.

Write a function `katz_centrality` that takes a graph `G`, number of iterations `n`, attenuation factor `beta` and returns a np.array of centrality values.

<!-- grade_id: cell-c8319e7048861c63, points: 0, grade: false, solution: true -->
```python
def katz_centrality(G, n, beta):
    '''Calculate Katz centrality'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-1ebe7a8d0b3a8b75, points: 1, grade: true, solution: false -->
```python
'''Check the restriction on beta and networkx implementation'''
beta = 1 / 3
try:
    katz = katz_centrality(G, 1, beta)
except:
    pass
else:
    raise AssertionError("Did not raise an exception")
beta = 1 / 5
assert 0.9999 < np.corrcoef(katz_centrality(G, 1, beta), c['degree'])[0][1] <= 1
nxkatz = np.array([val for val in nx.katz_centrality(G, beta).values()])
assert 0.9999 < np.corrcoef(katz_centrality(G, 20, beta), nxkatz)[0][1] <= 1
```

Let us draw the graph with Katz centrality visualization.

```python
katz_c = katz_centrality(subG, 20, 1/4)
katz_c = katz_c / katz_c.max()
plt.figure(figsize=(8, 6))
nx.draw(subG, 
        pos,
        width=0.5,
        linewidths=0.5,
        edgecolors='black',
        cmap=plt.cm.hot,
        node_size=katz_c*300,
        node_color=katz_c)
plt.show()
```

Here we can see how Katz centrality converges by the number of iterations.

```python
plt.figure(figsize=(7, 5))
convergence = []
old_katz = katz_centrality(G, 1, 0.2)
step_space = np.arange(2, 21)
for i in step_space:
    katz = katz_centrality(G, i, 0.2)
    convergence.append(np.linalg.norm(katz - old_katz))
    old_katz = katz.copy()
plt.plot(step_space, convergence)
plt.ylabel('Difference, Euclidean norm')
plt.xlabel('Step')
plt.title('Convergence of Katz centrality')
plt.show()
```

### Task 6. Eigenvector centrality (0.5 point)

Let us implement eigenvector centrality that can be interpreted as influence of a node. A centrality value of a node takes into account centrality values of its neighbors, that is

$$v_i = \frac{1}{\lambda}\sum_i A_{ij} v_j$$

or in the matrix form

$$Av = \lambda v$$

where $\lambda$ is the largest eigenvalue.

Write a function `eigenvector_centrality` that takes an adjacency matrix `A` and returns a np.array with centrality values.

<!-- grade_id: cell-8dc712b03819afba, points: 0, grade: false, solution: true -->
```python
def eigenvector_centrality(A):
    '''Calculate eigenvector centrality'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-15a0fcb2c488608e, points: 0.5, grade: true, solution: false -->
```python
'''Check some random values'''
eig_c = eigenvector_centrality(nx.to_numpy_array(subG))
assert eig_c.shape == (57, )
assert eig_c[0].round(4) == 0.0244
assert eig_c[-1].round(4) == 0.0701
assert eig_c[25].round(4) == 0.1198
```

Let us draw the graph with eigenvector centrality visualization.

```python
plt.figure(figsize=(8, 6))
nx.draw(subG, 
        pos,
        width=0.5,
        linewidths=0.5,
        edgecolors='black',
        cmap=plt.cm.hot,
        node_size=eig_c/eig_c.max()*400,
        node_color=eig_c)
plt.show()
```

### Task 7. Correlations (0.5 point)

Finally, let us see the relation between centralities. Calculate Pearson correlation coefficient and then we will draw scatter plots on pairs of centralities.

Write a function `pearson_correlation` that takes two np.arrays and returns Person correlation coefficient.

<!-- grade_id: cell-b5e6459a752a958b, points: 0, grade: false, solution: true -->
```python
def pearson_correlation(i_vals, j_vals):
    '''Calculate Pearson correlation coefficient'''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-2689d32a2c28cd6f, points: 0.5, grade: true, solution: false -->
```python
'''Check some coefficients'''
centvals = [
    ('Degree', list(nx.degree_centrality(G).values())),
    ('Closeness', list(nx.closeness_centrality(G).values())),
    ('Betweenness', list(nx.betweenness_centrality(G).values())),
    ('Katz', katz_centrality(G, 20, 1/5)), 
    ('Eigenvector', eigenvector_centrality(nx.to_numpy_array(G)))
]
assert pearson_correlation(centvals[0][1], centvals[0][1]) == 1
assert round(pearson_correlation(centvals[4][1], centvals[2][1]), 4) == 0.3358
assert round(pearson_correlation(centvals[4][1], centvals[3][1]), 4) == 0.9169
assert round(pearson_correlation(centvals[0][1], centvals[1][1]), 4) == 0.7211
assert round(pearson_correlation(centvals[3][1], centvals[4][1]), 4) == 0.9169
```

```python
plt.figure(figsize=(2*5, 5*5))
plt.subplots_adjust(hspace=0.4, wspace=0.4)
k = 1
for i in range(len(centvals)):
    for j in range(i + 1, len(centvals)):
        i_label, i_vals = centvals[i]
        j_label, j_vals = centvals[j]
        correlation = pearson_correlation(i_vals, j_vals)
        plt.subplot(5, 2, k)
        plt.scatter(i_vals, j_vals, s=5)
        plt.title('Correlation: {:.4f}'.format(correlation))
        plt.xlabel(i_label)
        plt.ylabel(j_label)
        k += 1
```

```python

```